# Ordinary Differential Equations Exercise 1

## Imports

```python
%matplotlib inline
import matplotlib.pyplot as plt
import numpy as np
import seaborn as sns
from scipy.integrate import odeint
from IPython.html.widgets import interact, fixed
```

## Euler's method

[Euler's method](http://en.wikipedia.org/wiki/Euler_method) is the simplest numerical approach for solving a first order ODE numerically. Given the differential equation

$$ \frac{dy}{dx} = f(y(x), x) $$

with the initial condition:

$$ y(x_0)=y_0 $$

Euler's method performs updates using the equations:

$$ y_{n+1} = y_n + h f(y_n,x_n) $$

$$ h = x_{n+1} - x_n $$

Write a function `solve_euler` that implements the Euler method for a 1d ODE and follows the specification described in the docstring:

<!-- grade_id: None, points: 0, grade: false, solution: true -->
```python
def solve_euler(derivs, y0, x):
    """Solve a 1d ODE using Euler's method.
    
    Parameters
    ----------
    derivs : function
        The derivative of the diff-eq with the signature deriv(y,x) where
        y and x are floats.
    y0 : float
        The initial condition y[0] = y(x[0]).
    x : np.ndarray, list, tuple
        The array of times at which of solve the diff-eq.
    
    Returns
    -------
    y : np.ndarray
        Array of solutions y[i] = y(x[i])
    """
       # YOUR CODE HERE
    y=np.array(x)
    y[0]=y0
    h=x[1]-x[0]
    for n in range(len(x)-1):
        y[n+1]=y[n]+h*derivs(y,x)
    return y

    
```

<!-- grade_id: odesex01a, points: 2, grade: true, solution: false -->
```python
assert np.allclose(solve_euler(lambda y, x: 1, 0, [0,1,2]), [0,1,2])
```

The [midpoint method]() is another numerical method for solving the above differential equation. In general it is more accurate than the Euler method. It uses the update equation:

$$ y_{n+1} = y_n + h f\left(y_n+\frac{h}{2}f(y_n,x_n),x_n+\frac{h}{2}\right) $$

Write a function `solve_midpoint` that implements the midpoint method for a 1d ODE and follows the specification described in the docstring:

<!-- grade_id: None, points: 0, grade: false, solution: true -->
```python
def solve_midpoint(derivs, y0, x):
    """Solve a 1d ODE using the Midpoint method.
    
    Parameters
    ----------
    derivs : function
        The derivative of the diff-eq with the signature deriv(y,x) where y
        and x are floats.
    y0 : float
        The initial condition y[0] = y(x[0]).
    x : np.ndarray, list, tuple
        The array of times at which of solve the diff-eq.
    
    Returns
    -------
    y : np.ndarray
        Array of solutions y[i] = y(x[i])
    """
    # YOUR CODE HERE
    y=np.array(x)
    y[0]=y0
    h=x[1]-x[0]
    for n in range(len(x)-1):
        y[n+1]=y[n]+h*derivs(y+0.5*h*derivs(y,x),x+0.5*h)
    return y
```

<!-- grade_id: odesex01b, points: 2, grade: true, solution: false -->
```python
assert np.allclose(solve_euler(lambda y, x: 1, 0, [0,1,2]), [0,1,2])
```

You are now going to solve the following differential equation:

$$
\frac{dy}{dx} = x + 2y
$$

which has the analytical solution:

$$
y(x) = 0.25 e^{2x} - 0.5 x - 0.25
$$

First, write a `solve_exact` function that compute the exact solution and follows the specification described in the docstring:

<!-- grade_id: , points: 0, grade: false, solution: true -->
```python
def solve_exact(x):
    """compute the exact solution to dy/dx = x + 2y.
    
    Parameters
    ----------
    x : np.ndarray
        Array of x values to compute the solution at.
    
    Returns
    -------
    y : np.ndarray
        Array of solutions at y[i] = y(x[i]).
    """
    # YOUR CODE HERE
    y=np.array(x)
    dy= x+2*y
    y=0.25*np.exp(2*x)-0.5*x-0.25
    return y
```

<!-- grade_id: odesex01c, points: 2, grade: true, solution: false -->
```python
assert np.allclose(solve_exact(np.array([0,1,2])),np.array([0., 1.09726402, 12.39953751]))
```

In the following cell you are going to solve the above ODE using four different algorithms:

1. Euler's method
2. Midpoint method
3. `odeint`
4. Exact

Here are the details:

* Generate an array of x values with $N=11$ points over the interval $[0,1]$ ($h=0.1$).
* Define the `derivs` function for the above differential equation.
* Using the `solve_euler`, `solve_midpoint`, `odeint` and `solve_exact` functions to compute
  the solutions using the 4 approaches.

Visualize the solutions on a sigle figure with two subplots:

1. Plot the $y(x)$ versus $x$ for each of the 4 approaches.
2. Plot $\left|y(x)-y_{exact}(x)\right|$ versus $x$ for each of the 3 numerical approaches.

Your visualization should have legends, labeled axes, titles and be customized for beauty and effectiveness.

While your final plot will use $N=10$ points, first try making $N$ larger and smaller to see how that affects the errors of the different approaches.

<!-- grade_id: None, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
x=np.linspace(0,1,11)
ic=np.array([0,0])

def derivs(yvec,x):
    y=yvec[0]
    dy= x+2*y
    return np.array([dy])


soln= odeint(derivs,
            ic,
            x)

euler=solve_euler(derivs, 0, x)
# middy=solve_midpoint(derivs, 0, x)
exact=solve_exact(x)
plt.figure(figsize=(9,6))
plt.legend()
plt.plot(x,soln[:,0])
plt.plot(x,exact)
plt.title('Comparison')
plt.ylabel('Y-Axis')
plt.xlabel("X-Axis")


```

Output:
```
<ipython-input-118-67f53d2ea25e> in <module>()
     13             x)
     14 
---> 15 euler=solve_euler(derivs, 0, x)
     16 # middy=solve_midpoint(derivs, 0, x)
     17 exact=solve_exact(x)

<ipython-input-111-be147e157977> in solve_euler(derivs, y0, x)
     22     h=x[1]-x[0]
     23     for n in range(len(x)-1):
---> 24         y[n+1]=y[n]+h*derivs(y,x)
     25     return y
     26 

ValueError: setting an array element with a sequence.
```

<!-- grade_id: odesex01d, points: 4, grade: true, solution: false -->
```python
assert True # leave this for grading the plots
```

```python

```