# Project Euler: Problem 52

https://projecteuler.net/problem=52

It can be seen that the number, $125874$, and its double, $251748$, contain exactly the same digits, but in a different order.

Find the smallest positive integer, $x$, such that $2x$, $3x$, $4x$, $5x$, and $6x$, contain the same digits.

First, write a function `same_digits(x,y)` that returns `True` if two integers `x` and `y` have the exact same set of digits and multiplicities and `False` if they have different digits.

<!-- grade_id: None, points: 0, grade: false, solution: true -->
```python
import numpy as np
i=0
j=0
def same_digits(x, y):
    """Do the integers x and y have the same digits, regardless of order."""
    # YOUR CODE HERE
    stringx=str(x)
    stringy=str(y)
    a=np.zeros((len(stringx)))
    b=np.zeros((len(stringy)))

    for i in range(len(stringx)):
        a[i]=stringx[i]
    for j in range(len(stringy)):
        b[j]=stringy[j]
    w=np.sort(a)
    z=np.sort(b)
    if len(stringy)!=len(stringx):
        return False
    for k in range (len(stringx)):
        if w[k]!=z[k]:
            return False
        else:
            return True
#     if w.all==z.all:
#         return True
#     else:
#         return False
#     print(w)
#     print(z)
#     return

same_digits(12273,2321)

    
# x=12345
# string=str(x)
# string[0]
# len(string)
# a=np.zeros((len(string)))
# a[0]=5
# a
# for i in range(len(string)):
#     a[i]=string[i]

# print(a)

    
    
# numbers=np.array(list(x),dtype=np.int)
# numbers

# ref={}
# ref[1]='hi'
# ref
```

Output:
```
False
False
False
False
```

<!-- grade_id: projecteuler52a, points: 4, grade: true, solution: false -->
```python
assert same_digits('132', '321')
assert not same_digits('123', '3')
assert not same_digits('456', '0987654321')
```

Now use the `same_digits` function to solve this Euler problem. As you work on this problem, be careful to debug and test your code on small integers before trying it on the full search.

<!-- grade_id: None, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
i=2
def check(x):
    for i in range (x):
        if same_digits(i,2*i):#and same_digits(i,3*i) and same_digits(i,4*i) and same_digits(i,5*i) and same_digits(i,6*i)    
            return i
y=check(100000000)
y
    

```

Output:
```
0
0
0
0
```

<!-- grade_id: projecteuler52b, points: 6, grade: true, solution: false -->
```python
assert True # leave this cell to grade the solution
```

```python

```

```python

```

```python

```

```python

```