# Project Euler: Problem 8

https://projecteuler.net/problem=8

The four adjacent digits in the 1000-digit number that have the greatest product are 9 × 9 × 8 × 9 = 5832.

(see the number below)

Find the thirteen adjacent digits in the 1000-digit number that have the greatest product. What is the value of this product?

Use NumPy for this computation

```python
import numpy as np
```

```python
d1000 = """
73167176531330624919225119674426574742355349194934
96983520312774506326239578318016984801869478851843
85861560789112949495459501737958331952853208805511
12540698747158523863050715693290963295227443043557
66896648950445244523161731856403098711121722383113
62229893423380308135336276614282806444486645238749
30358907296290491560440772390713810515859307960866
70172427121883998797908792274921901699720888093776
65727333001053367881220235421809751254540594752243
52584907711670556013604839586446706324415722155397
53697817977846174064955149290862569321978468622482
83972241375657056057490261407972968652414535100474
82166370484403199890008895243450658541227588666881
16427171479924442928230863465674813919123162824586
17866458359124566529476545682848912883142607690042
24219022671055626321111109370544217506941658960408
07198403850962455444362981230987879927244284909188
84580156166097919133875499200524063689912560717606
05886116467109405077541002256983155200055935729725
71636269561882670428252483600823257530420752963450
"""
```

```python
#solution using Numpy
a=np.zeros((20,50))
i=1
j=0
k=0
big=1
biggest=1
count=0
while i<1020:
    a[j,k]=d1000[i]
    if i%(50+(count*51))==0:
        j=j+1 #move down one column
        i=i+1 #index the string again
        k=-1  #reset at the first column
        count=count+1 #be ready for next row change    
    i=i+1
    k=k+1
# print(a)
    
x=0
y=0
for x in range(20):    
    for y in range (38):
        b=(a[x,y:y+13])
#         print(b)
        mult=b.cumprod()
#         l=[int(x) for x in mult]    
#         print(l)
        big=mult[12]
        if big>biggest:
            biggest=big    
print(biggest)




# b=(a[j,k:k+4])
# mult=b.cumprod()
# big=mult[3]
# if big>biggest:
#         biggest=big       
# print(big)


```

Output:
```
5377010688.0
```

<!-- grade_id: None, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
#working code for multiplying 4 and 13 numbers not using numpy

# i=2
# big=1
# count =0
# while i<1000:
#     if i%(48+(count*51))==0:
#         i=i+4
#         count= count+1
#     a=d1000[i:i+4]
#     num =1
#     for b in a:
#         num=num*int(b)
#     if num>big:
#         big =num
#     i=i+1
# print (big)

i=2
big=1
count =0
while i<1000:#run over the length of the number
    if i%(39+(count*51))==0: #if index reaches a non integer move to next line
        i=i+13
        count= count+1
    a=d1000[i:i+13]#create array 13 digits long
    num =1
    for b in a:#multiply all 13 digits by each other
        num=num*int(b)
    if num>big: #find the biggest multiplication
        big =num
    i=i+1#scoot array down one index along rows
print (big)
   

```

Output:
```
5377010688
```

<!-- grade_id: projecteuler8, points: 10, grade: true, solution: false -->
```python
assert True # leave this for grading
```